# Design Decisions with Python Functions



```python
from nose.tools import assert_true, assert_equal, assert_false, assert_almost_equal, assert_raises
```

# Two primary reasons for defining functions:
1. Code reuse: 
    * Write and debug code once. Then I can use this same (correct) code many times
    * This makes upkeep/modifications simpler. When I think of an improved way of implementing something I only need to change it in one location.

2. Procedural Decomposition:
    * A function should do one thing, not multiple things.
    * This can become a matter of style
    

## Function Style According to Mark Thomason
![Mark Thomason](./mark_thomason.jpg)

## The entire function should be visible on your screen within your editor.
## If your function doesn't fit on your screen, get a bigger screen

### What are the implications of these heuristic?


## Exercise: Define a function to get a positive integer from a user.
### Requirements
1. Use an infinite while loop
1. Use the input function
1. Keep prompting the user for input until a valid positive integer is provided



```python
def get_pos_integer(prompt="Enter a positive integer"):
    while True:
        ???

```

## Does this function do one thing?
## Could we break it into smaller pieces?

### Write a function to test whether a number is positive

<!-- grade_id: cell-b17130446493e1e7, points: 0, grade: false, solution: true -->
```python
def ispositive(x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-163d61c8d4528ae4, points: 0, grade: true, solution: false -->
```python
assert_true(ispositive(5))
assert_false(ispositive(-1))
assert_false(ispositive(0))
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-1-fc1014b3334c> in <module>()
----> 1 assert_true(ispositive(5))
      2 assert_false(ispositive(-1))
      3 assert_false(ispositive(0))

NameError: name 'assert_true' is not defined
```

## Exercise

Following the same style as `get_pos_integer`, write a function `get_value`. `get_value` takes as arguments:

1. A positional argument `converter` that is a function that takes as input a string and returns the desired value
1. A positional argument `tester` that is a function that takes as input a value and returns `True` or `False` depending on whether a desired condition is satisfied.
1. A keyword argument `prompt` that is the prompt to use with `input`.

Test the function with `getint` and `ispositive`.

<!-- grade_id: cell-cebe2a0d20841727, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
get_value(getint, ispositive)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-2-8e5c70177eab> in <module>()
----> 1 get_value(getint, ispositive)

NameError: name 'get_value' is not defined
```

## Exercise

Modify `test_ascending` to take in a sequence and test if the elements in `values` are in ascending order.

**Hint:** Use the `all` function.

<!-- grade_id: cell-40b8ef1cd8e54c0b, points: 0, grade: false, solution: true -->
```python
def test_ascending(values):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-519436fcb56d35ad, points: 0, grade: true, solution: false -->
```python
assert_true(test_ascending(("Argos", "Helios", "Zeus")))
assert_false(test_ascending(("Argos", "Zeus", "Helios")))
assert_false(test_ascending(("argos", "Helios", "Zeus")))

```

```python
get_value(get_three_words, 
          test_ascending, 
          prompt="enter three words in ascending alphabetical order separated by spaces: ")

```